<a href="https://colab.research.google.com/github/gleetx/bootcamp/blob/main/Module%200/Session%203/m0s3nb4_mpe_3%20GL.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Problem 3

This problem has a single exercise worth a total of ten (10) points.

**Exercise 0** (10 points). Define a function, `UniqueCharacters(s)`, that given a string `s` returns a tuple of two elements: the first element is the number of unique alphabetic characters in the string and the second is the number of unique digits (base-10) in the string.

For example, the string `'ewwwffioj122434'` should output the following tuple: `(6, 4)`. The `6` occurs because there are 6 unique letters (`'e'`, `'w'`, `'f'`, `'i'`, `'o'`, and `'j'`) and `4` because there are 4 unique decimal digits (`'1'`,  `'2'`, `'3'`, `'4'`). Special characters may appear in the string but do not count as a letter or number.

```python
def UniqueCharacters(s):
    ### YOUR CODE HERE

#     GOAL
#         Tuple of 2 elements:
#         the first element is the number of unique alphabetic characters in the string and
#         the second is the number of unique digits (base-10) in the string.

#     INPUT
#           string called s

#     STRATEGY
#             create letter and number list variables
#             loop over the input string
#             detect if the character is a letter or number
#                 https://duckduckgo.com/?q=python+is+character+alphabetic&t=chromentp&ia=web

#             append the letter or number to the correct list variable
#             after the loop take the unique values of the two variables
#                     https://duckduckgo.com/?q=python+unique+values+in+a+list&t=chromentp&ia=web

#             create tuple of the lengths of each variable
#             return the tuple
#     SOLUTION

    lets = []
    nums = []

    for char in s:
        if char.isalpha():
            lets.append(char)
        if char.isdigit():
            nums.append(char)

    set_lets = set(lets)
    set_nums = set(nums)

    ret_var = (len(set_lets),len(set_nums))

    return ret_var


# test case
# UniqueCharacters('ewwwffioj122434')
```

```python
### doing the elif method vs else method below so we don't have to account for non alpha or non digit characters.
def UniqueCharacters2(s):
  unique_set = set(s)
  print(unique_set)

  alpha_count = 0
  digit_count = 0

  for character in unique_set:
    if character.isalpha():
      alpha_count += 1
    elif character.isdigit():
      digit_count += 1

  return(alpha_count,digit_count)





# test case 2
UniqueCharacters2('ehruebuv837483dhf')

```

Output:
```
{'7', 'b', '4', 'u', '3', 'f', 'd', 'r', 'v', 'h', 'e', '8'}
(8, 4)
```

<!-- grade_id: UniqueCharacters, points: 10, grade: true, solution: false -->
```python
# Test Cell: 'UniqueCharacters' (10 points)
assert UniqueCharacters('abc123') == (3,3)
assert UniqueCharacters('abc') == (3,0)
assert UniqueCharacters('aaa') == (1,0)
assert UniqueCharacters('aaa111') == (1,1)
assert UniqueCharacters('111') == (0,1)
assert UniqueCharacters('123') == (0,3)
assert UniqueCharacters('') == (0,0)
assert UniqueCharacters('///;;;...,????!!!!!###$$$%%%') == (0,0)
assert UniqueCharacters('//23bd') == (2,2)
assert UniqueCharacters('b2b3n4s4d9') == (4,4)
assert UniqueCharacters('9090909p0y90p90y90') == (2,2)
assert UniqueCharacters('ieowjfiojfioj2342io4ji') == (6,3)
assert UniqueCharacters('ewwwffioj122434') == (6,4)
assert UniqueCharacters('dwdj2ru3jf894jgf.,/23,4./3ei8fj2389ej89/,.,./2dd32je98dej89ij') == (9,5)
print("\n(Passed!)")
```

Output:
```
(Passed!)
```

**Fin!** You've reached the end of this problem. Don't forget to restart the
kernel and run the entire notebook from top-to-bottom to make sure you did
everything correctly. If that is working, try submitting this problem. (Recall
that you *must* submit and pass the autograder to get credit for your work!)